```python
from grading_tools import check, auto_marking_message
import cvxopt
import numpy as np
from math import *
import matplotlib.pyplot as plt
```

# Exercises

## Exercise

If you invest $\$2500$ in Google stock and it grows to have a value of $\$3500$, what is your return?  



<!-- grade_id: 4778b9, points: 0, grade: false, solution: true -->
```python
(3500-2500)/2500
```

Output:
```
0.4
```

<!-- grade_id: 56c4ed, points: 1, grade: true, solution: false -->
```python
check('56c4ed',_)
auto_marking_message()
```

Output:
```
Auto marking message: âœ¨ Correct
```

## Exercise

Plot a chart of the historic weekly return of the first stock in the `ukx.xlsx` file. It should be very spiky and look a lot a graph of white noise.

How would you generate a sequence of normally distributed random numbers with the same mean and standard deviation as the returns of this first stock? Generate a plot of such random numbers that can be compared with the plot of historic weekly returns.

Can you visually detect any differences between the plots?

<!-- grade_id: 60c910, points: 0, grade: false, solution: true -->
```python
# So you don't have to keep retyping the same
# code, I've created two handy functions read_return_history and solve_markowitz
# in the file markowitzlib.py which you can use. 
from markowitzlib import read_return_history

# --- Get the matrix of returns ---
return_history = read_return_history()

# --- Take first stock's returns ---
first_stock_returns = return_history[:, 0]

# --- Plot the historic weekly returns ---
plt.figure(figsize=(10, 4))
plt.plot(first_stock_returns, label="Historic Weekly Returns")
plt.xlabel("Week")
plt.ylabel("Return")
plt.title("Historic Weekly Return of First Stock")
plt.legend()
plt.show()

# --- Generate a synthetic normal series with same mean/std ---
mu = np.mean(first_stock_returns)
sigma = np.std(first_stock_returns)
synthetic_returns = np.random.normal(mu, sigma, len(first_stock_returns))

# --- Plot the synthetic returns ---
plt.figure(figsize=(10, 4))
plt.plot(synthetic_returns, color="orange", label="Synthetic (Normal) Returns")
plt.xlabel("Week")
plt.ylabel("Return")
plt.title("Synthetic Normally Distributed Returns (same Î¼, Ïƒ)")
plt.legend()
plt.show()

# --- Print comparison stats ---
print(f"Historic mean = {mu:.6f}, std = {sigma:.6f}")
print(f"Synthetic mean = {np.mean(synthetic_returns):.6f}, std = {np.std(synthetic_returns):.6f}")
```

Output:
```
File c:\Users\peiya\AppData\Local\Programs\Python\Python313\Lib\site-packages\pandas\io\excel\_openpyxl.py:552, in OpenpyxlReader.__init__(self, filepath_or_buffer, storage_options, engine_kwargs)
    534 @doc(storage_options=_shared_docs["storage_options"])
    535 def __init__(
    536     self,
   (...)    539     engine_kwargs: dict | None = None,
    540 ) -> None:
    541     """
    542     Reader using openpyxl engine.
    543 
   (...)    550         Arbitrary keyword arguments passed to excel engine.
    551     """
--> 552     import_optional_dependency("openpyxl")
    553     super().__init__(
    554         filepath_or_buffer,
    555         storage_options=storage_options,
    556         engine_kwargs=engine_kwargs,
    557     )

File c:\Users\peiya\AppData\Local\Programs\Python\Python313\Lib\site-packages\pandas\compat\_optional.py:138, in import_optional_dependency(name, extra, errors, min_version)
    136 except ImportError:
    137     if errors == "raise":
--> 13
```

## Exercise

The file `goog-options.xlsx` contains data for option prices on Google stock taken on one particular day. Use the "explorer" tab on Cocalc to view all the files in this directory which includes this goog\-options.xslx file so you can download it to view it.

Each row in the main table represents the price of a different call option. The first column contains an identifier for the specific call option contract. The word _security_ is used to refer to a general financial product whether it is a stock, bond, option \(or any other kind of financial product\), so this identifier is called a _security identifier_.

The second column contains the *last price* of the option and the third price is the *mid price*
There are two main prices quoted for any financial product on the market. The ask price is how much you must pay to buy a product, the bid price is how much you will receive if you sell the product. The mid price is the average of the two. The last price is the last price traded, which may be either the bid or the ask.

The final column contains the strike price of each option.

Plot a graph of the mid price against the strike price.

So I can test your code, store the sum of all the mid prices in a variable called `sum_mids`.



<!-- grade_id: 732b6f, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import matplotlib.pyplot as plt


# 1) Load
df = pd.read_excel("goog-options.xlsx")

# 2) Infer/normalize columns â†’ security, last, mid, strike
lower_cols = [str(c).strip().lower() for c in df.columns]
name_map = {}
for i, c in enumerate(lower_cols):
    if "sec" in c:      name_map[i] = "security"
    elif "last" in c:   name_map[i] = "last"
    elif "mid" in c:    name_map[i] = "mid"
    elif "strike" in c: name_map[i] = "strike"

# Fallback to positional order if needed: [security, last, mid, strike]
if set(name_map.values()) != {"security", "last", "mid", "strike"}:
    if df.shape[1] < 4:
        raise ValueError("Expected at least 4 columns: security, last, mid, strike.")
    name_map = {0: "security", 1: "last", 2: "mid", 3: "strike"}

# Rename
df = df.rename(columns={df.columns[i]: new for i, new in name_map.items()})

# 3) Keep only needed columns, coerce to numeric
df = df[["security", "last", "mid", "strike"]].copy()
df["mid"] = pd.to_numeric(df["mid"], errors="coerce")
df["strike"] = pd.to_numeric(df["strike"], errors="coerce")

# 4) Optional: drop rows missing mid/strike and sort by strike (helps plotting)
df = df.dropna(subset=["mid", "strike"]).sort_values("strike")

# 5) Sum of all mid prices (as required)
sum_mids = df["mid"].sum()

# 6) Plot mid vs strike
plt.figure()
plt.plot(df["strike"], df["mid"], marker="o", linestyle="-")
plt.xlabel("Strike price")
plt.ylabel("Mid price")
plt.title("GOOG Call Options: Mid vs Strike")
plt.grid(True)
plt.tight_layout()
plt.show()

print("sum_mids =", sum_mids)
```

Output:
```
File c:\Users\peiya\AppData\Local\Programs\Python\Python313\Lib\site-packages\pandas\io\excel\_openpyxl.py:552, in OpenpyxlReader.__init__(self, filepath_or_buffer, storage_options, engine_kwargs)
    534 @doc(storage_options=_shared_docs["storage_options"])
    535 def __init__(
    536     self,
   (...)    539     engine_kwargs: dict | None = None,
    540 ) -> None:
    541     """
    542     Reader using openpyxl engine.
    543 
   (...)    550         Arbitrary keyword arguments passed to excel engine.
    551     """
--> 552     import_optional_dependency("openpyxl")
    553     super().__init__(
    554         filepath_or_buffer,
    555         storage_options=storage_options,
    556         engine_kwargs=engine_kwargs,
    557     )

File c:\Users\peiya\AppData\Local\Programs\Python\Python313\Lib\site-packages\pandas\compat\_optional.py:138, in import_optional_dependency(name, extra, errors, min_version)
    136 except ImportError:
    137     if errors == "raise":
--> 13
```

<!-- grade_id: 7420b9, points: 1, grade: true, solution: false -->
```python
check('7420b',sum_mids)
auto_marking_message()
```

Output:
```
Auto marking message: ðŸ† Correct
```

## Exercise

Suppose that short-selling is not allowed, so all the components of the weight vector must be positive. Write down a version of the optimization problem where short-selling is not allowed. Write a Python function that can solve this problem given a covariance matrix $\Sigma$ and mean vector $\mu$. This should work exactly like the `solve_markowitz` function in the handout except that it should be able to solve
the problem when short-selling is allowed and when short-selling is not allowed.

If short-selling is not allowed, then not all return values can actually be achieved, so sometimes your program will fail saying that the optimization is not feasible. However, I found that for the $\Sigma$ and $\mu$ used in the previous video that you can find solutions for expected returns in the range from 0 to 0.007.

If your code is correct, then the tests will plot a graph showing how the efficient frontier changes when short-selling is allowed
or not allowed.

<!-- grade_id: 39bbd4, points: 0, grade: false, solution: true -->
```python
# The next line stops the CVX optimiser filling your notebook with
# output: it produces a lot of annoying print out if you run it
# to solve the problem with short-selling prohibited.
cvxopt.solvers.options['show_progress'] = False

def solve_markowitz( sigma, mu, R, allow_short_selling = True ):
    # YOUR CODE HERE
    raise NotImplementedError()

from markowitzlib import read_return_history as rrh
return_history = rrh()
mu = np.mean( return_history, axis = 0)
print(mu[0])
```

Output:
```
File c:\Users\peiya\AppData\Local\Programs\Python\Python313\Lib\site-packages\pandas\io\excel\_openpyxl.py:552, in OpenpyxlReader.__init__(self, filepath_or_buffer, storage_options, engine_kwargs)
    534 @doc(storage_options=_shared_docs["storage_options"])
    535 def __init__(
    536     self,
   (...)    539     engine_kwargs: dict | None = None,
    540 ) -> None:
    541     """
    542     Reader using openpyxl engine.
    543 
   (...)    550         Arbitrary keyword arguments passed to excel engine.
    551     """
--> 552     import_optional_dependency("openpyxl")
    553     super().__init__(
    554         filepath_or_buffer,
    555         storage_options=storage_options,
    556         engine_kwargs=engine_kwargs,
    557     )

File c:\Users\peiya\AppData\Local\Programs\Python\Python313\Lib\site-packages\pandas\compat\_optional.py:138, in import_optional_dependency(name, extra, errors, min_version)
    136 except ImportError:
    137     if errors == "raise":
--> 13
```

<!-- grade_id: 1cf82d, points: 1, grade: true, solution: false -->
```python
from markowitzlib import read_return_history as rrh
return_history = rrh()

mu = np.mean( return_history, axis=0 )
sigma = np.cov( return_history, rowvar=False )
res = solve_markowitz(sigma, mu, 0.003, True)
check( '1cf82d.1', res['sd'], tol=0.000001 )
res = solve_markowitz(sigma, mu, 0.003, False)
check( '1cf82d.2', res['sd'], tol =0.000001 )

for allow_short_selling in [True, False]:
    mu = np.mean( return_history, axis=0 )
    sigma = np.cov( return_history, rowvar=False )
    returns = np.linspace(0,0.007,50)
    min_sds = np.zeros(len(returns))
    for i in range(0,len(returns)):
        ret = solve_markowitz(sigma, mu, returns[i], allow_short_selling)
        min_sds[i] = ret['sd']

    ax = plt.gca()
    if allow_short_selling:
        l = 'Short-selling allowed'
    else:
        l = 'Short-selling prohibited'
    ax.plot(min_sds,returns, label=l);
    
ax.legend();
auto_marking_message()
```

## Exercise

Suppose that the market contains $n$ stocks. Suppose that the returns of the stocks over the time period $0$ to $T$ are given by independent identically distributed random variables
$R_1$, $R_2$, ..., $R_n$ each with mean $\mu$ and standard deviation $\sigma$. Find the weights of a portfolio whose return is $\mu$
and whose standard deviation is
$$\frac{1}{\sqrt{n}}\sigma.$$

**Hint:** What is the most diversified possible investment across these stocks?



<!-- grade_id: cbd518, points: 0, grade: false, solution: true -->
```python
def find_weights( n ):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: a95832, points: 1, grade: true, solution: false -->
```python
n = 7
w = find_weights(n)
assert len(w)==7
check('cbd518.1', w.transpose() @ w)
check('cbd518.2', sum(w))
auto_marking_message()
```

## Exercise

Consider an $n$\-asset portfolio as defined in lectures. Using our expression for the weights for the minimum variance portfolio for a given return $R$, it can be shown that the efficient frontier is 
$$
\sigma^2_p = \frac{AR^2 - 2BR + C}{AC-B^2}.
$$
What is the global minimum variance portfolio? Provide its expected return and weights as well as the variance.



## Exercise

Consider a portfolio of 3 assets with means' ${\bf\mu}=(0.1, 0.05,0.15)$ and covariance matrix $\Sigma = I_3$, the $3\times 3$ identity matrix.
Using the method of Lagrange multipliers, determine the weights of the minimum variance portfolio that achieve return $R$.

Sketch the efficient frontier in this case, marking the global minimum variance point and $\sigma$ intercept.

